# Sentiment Analysis with NLTK

## Python setup


We need the following modules:

- NLTK
- Pandas
- Altair

```python
import nltk

# we suppress some unimportant warnings
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
```

## Data

### Data import

```python
import pandas as pd

# Import some prepared twitter data from cnn breaking news
df = pd.read_csv("https://raw.githubusercontent.com/kirenz/datasets/master/sentiment-cnn.csv")

df.head(3)
```

Output:
```
text  \
0  the body of missing princeton university stude...   
1  uk prime minister liz truss quits after a disa...   
2  trump weighs letting federal agents return to ...   

                  created_at  \
0  2022-10-20 19:58:17+00:00   
1  2022-10-20 12:37:10+00:00   
2  2022-10-19 23:03:37+00:00   

                                          text_token  \
0  ['the', 'body', 'of', 'missing', 'princeton', ...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weighs', 'letting', 'federal', 'age...   

                                        text_token_s  \
0  ['body', 'missing', 'princeton', 'university',...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weighs', 'letting', 'federal', 'age...   

                                             text_si  \
0  body missing princeton university student misr...   
1  prime minister liz truss quits disastrous six ...   
2  trump weighs letting federal agents return mar...   

                   
```

### Data corrections

```python
df['created_at'] = pd.to_datetime(df['created_at'])

df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 22 entries, 0 to 21
Data columns (total 6 columns):
 #   Column        Non-Null Count  Dtype              
---  ------        --------------  -----              
 0   text          22 non-null     object             
 1   created_at    22 non-null     datetime64[ns, UTC]
 2   text_token    22 non-null     object             
 3   text_token_s  22 non-null     object             
 4   text_si       22 non-null     object             
 5   text_sil      22 non-null     object             
dtypes: datetime64[ns, UTC](1), object(5)
memory usage: 1.2+ KB
```

## Analysis

### VADER lexicon

- NLTK provides a simple rule-based model for general sentiment analysis called VADER, which stands for "Valence Aware Dictionary and Sentiment Reasoner" (Hutto & Gilbert, 2014).

```python
nltk.download('vader_lexicon')
```

Output:
```
[nltk_data] Downloading package vader_lexicon to
[nltk_data]     /Users/jankirenz/nltk_data...
[nltk_data]   Package vader_lexicon is already up-to-date!
True
```

### Sentiment Intensity Analyzer

- Initialize an object of `SentimentIntensityAnalyzer` with name "analyzer":

```python
from nltk.sentiment import SentimentIntensityAnalyzer

analyzer = SentimentIntensityAnalyzer()
```

### Polarity scores

- Use the `polarity_scores` method:

```python
df['polarity'] = df['text_sil'].apply(lambda x: analyzer.polarity_scores(x))
```

```python
df.head(3)
```

Output:
```
text  \
0  the body of missing princeton university stude...   
1  uk prime minister liz truss quits after a disa...   
2  trump weighs letting federal agents return to ...   

                 created_at  \
0 2022-10-20 19:58:17+00:00   
1 2022-10-20 12:37:10+00:00   
2 2022-10-19 23:03:37+00:00   

                                          text_token  \
0  ['the', 'body', 'of', 'missing', 'princeton', ...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weighs', 'letting', 'federal', 'age...   

                                        text_token_s  \
0  ['body', 'missing', 'princeton', 'university',...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weighs', 'letting', 'federal', 'age...   

                                             text_si  \
0  body missing princeton university student misr...   
1  prime minister liz truss quits disastrous six ...   
2  trump weighs letting federal agents return mar...   

                       
```

### Transform data

```python
# Change data structure (we unnest the column polarity and add new columns)
df = pd.concat([df.drop(['polarity'], axis=1), df['polarity'].apply(pd.Series)], axis=1)
```

```python
df.head()
```

Output:
```
text  \
0  the body of missing princeton university stude...   
1  uk prime minister liz truss quits after a disa...   
2  trump weighs letting federal agents return to ...   
3  britain's home secretary suella braverman quit...   
4  russian president vladimir putin signs a decre...   

                 created_at  \
0 2022-10-20 19:58:17+00:00   
1 2022-10-20 12:37:10+00:00   
2 2022-10-19 23:03:37+00:00   
3 2022-10-19 16:17:29+00:00   
4 2022-10-19 12:02:15+00:00   

                                          text_token  \
0  ['the', 'body', 'of', 'missing', 'princeton', ...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weighs', 'letting', 'federal', 'age...   
3  ['britain', 's', 'home', 'secretary', 'suella'...   
4  ['russian', 'president', 'vladimir', 'putin', ...   

                                        text_token_s  \
0  ['body', 'missing', 'princeton', 'university',...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weigh
```

Create new variable called sentiment which contains the entries "neutral," "positive" or "negative" (depending on the compound score).


Hint:


---

```python
df['___'] = df['___'].___(___ x: '___' if ___ >___ else '___' if ___ else '___')
```

---

- Name the new variable `sentiment`
- Use variable `compound` as basis
- apply a lambda function to each row.
- The lambda function should write a name in a cell:
  - 'positive' `if x>0`
  - 'neutral' `if x==0`
  - 'negative' for all other cases (`else`)


<!-- grade_id: sentiment, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
df['sentiment'] = df['compound'].apply(lambda x: 'positive' if x >0 else 'neutral' if x==0 else 'negative')
### END SOLUTION
```

<!-- grade_id: c-sentiment, points: 4, grade: true, solution: false -->
```python
# check your code
assert df.iloc[0, 10] == 'negative'
```

```python
df.head()
```

Output:
```
text  \
0  the body of missing princeton university stude...   
1  uk prime minister liz truss quits after a disa...   
2  trump weighs letting federal agents return to ...   
3  britain's home secretary suella braverman quit...   
4  russian president vladimir putin signs a decre...   

                 created_at  \
0 2022-10-20 19:58:17+00:00   
1 2022-10-20 12:37:10+00:00   
2 2022-10-19 23:03:37+00:00   
3 2022-10-19 16:17:29+00:00   
4 2022-10-19 12:02:15+00:00   

                                          text_token  \
0  ['the', 'body', 'of', 'missing', 'princeton', ...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weighs', 'letting', 'federal', 'age...   
3  ['britain', 's', 'home', 'secretary', 'suella'...   
4  ['russian', 'president', 'vladimir', 'putin', ...   

                                        text_token_s  \
0  ['body', 'missing', 'princeton', 'university',...   
1  ['uk', 'prime', 'minister', 'liz', 'truss', 'q...   
2  ['trump', 'weigh
```

### Max and min sentiment

```python
# Tweet with highest positive sentiment
df[['text', 'compound', 'neg', 'neu', 'pos', 'sentiment']].loc[df['compound'].idxmax()]
```

Output:
```
text         bruce sutter, a cy young award-winning relief ...
compound                                                0.8834
neg                                                      0.126
neu                                                      0.386
pos                                                      0.488
sentiment                                             positive
Name: 16, dtype: object
```

```python
# Tweet with highest negative sentiment 
# ...seems to be a case of wrong classification because of the word "deficit"
df[['text', 'compound', 'neg', 'neu', 'pos', 'sentiment']].loc[df['compound'].idxmin()]
```

Output:
```
text         jury finds paul flores guilty of first-degree ...
compound                                               -0.9531
neg                                                       0.58
neu                                                      0.337
pos                                                      0.083
sentiment                                             negative
Name: 6, dtype: object
```

### Visualize data

```python
import altair as alt

# create data to change colors in Altair plot
domain = ['neutral', 'positive', 'negative']
range_=['#b2d8d8',"#008080", '#db3d13']


alt.Chart(df).mark_bar().encode(
    x=alt.X('count()', title=None),
    y=alt.Y('sentiment', sort="-x"),
    color= alt.Color('sentiment', legend=None, scale=alt.Scale(domain=domain, range=range_))
).properties(
    title="Sentiment analysis",
    width=400,
    height=150,
)
```

Output:
```
alt.Chart(...)
```

```python
# Function to add date variables to DataFrame.
def add_date_info(df):
  df['created_at'] = pd.to_datetime(df['created_at'], unit='ns')
  df['Year'] = pd.DatetimeIndex(df['created_at']).year
  df['Month'] = pd.DatetimeIndex(df['created_at']).month
  df['Day'] = pd.DatetimeIndex(df['created_at']).day
  df['DOY'] = pd.DatetimeIndex(df['created_at']).dayofyear
  df['Date'] = pd.DatetimeIndex(df['created_at']).date
  return df
```

```python
add_date_info(df)
```

Output:
```
text  \
0   the body of missing princeton university stude...   
1   uk prime minister liz truss quits after a disa...   
2   trump weighs letting federal agents return to ...   
3   britain's home secretary suella braverman quit...   
4   russian president vladimir putin signs a decre...   
5   rising food and housing costs drove uk inflati...   
6   jury finds paul flores guilty of first-degree ...   
7   russian expat who was the main source for infa...   
8   a pilot and passenger are dead after a plane c...   
9   ukraine's president zelensky says 30% of the c...   
10  president biden announces the formal launch of...   
11  a russian fighter jet has crashed into a resid...   
12  former president donald trump's company charge...   
13  steve bannon sentencing recommendation seeks s...   
14  blasts were heard in the ukrainian capital as ...   
15  authorities in central california have arreste...   
16  bruce sutter, a cy young award-winning relief ...   
17  judge throws out on
```

```python
# change format
df['Date'] = pd.to_datetime(df['Date'])
```

```python
alt.Chart(df).mark_area().encode(
   x=alt.X('Date', axis=alt.Axis(format='%e.%-m.')),
   y=alt.Y('count(sentiment)'),
   color=alt.Color('sentiment', scale=alt.Scale(domain=domain, range=range_))
)
```

Output:
```
alt.Chart(...)
```

```python
alt.Chart(df).mark_boxplot().encode(
    x=alt.X('sentiment'),
    y=alt.Y('compound'),
    color=alt.Color('sentiment', scale=alt.Scale(domain=domain, range=range_))
).properties(
    width=200,
    height=200
)
```

Output:
```
alt.Chart(...)
```

Literature:

[Hutto, C.J. & Gilbert, E.E. (2014). VADER: A Parsimonious Rule-based Model for
Sentiment Analysis of Social Media Text. Eighth International Conference on
Weblogs and Social Media (ICWSM-14). Ann Arbor, MI, June 2014.](https://ojs.aaai.org/index.php/ICWSM/article/view/14550)